# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## Working with environments

Working in Julia scripts the default setting is to add any packages we use to the default project at ~/.julia/environments/v1.0. 

It is however easy to create other, independent, projects. In order to create a new project, create a directory for it and then activate that directory to make it the "active project", which package operations manipulate:

`
shell> mkdir MyProject

shell> cd MyProject
/Users/kristoffer/MyProject

(v1.0) pkg> activate .

(MyProject) pkg> st
    Status `Project.toml`
    `

```julia
Pkg.installed()
```

Output:
```
┌ Warning: Pkg.installed() is deprecated
└ @ Pkg /Users/julia/buildbot/worker/package_macos64/build/usr/share/julia/stdlib/v1.6/Pkg/src/Pkg.jl:566
Dict{String, VersionNumber} with 16 entries:
  "DataVoyager"       => v"1.0.2"
  "CSVFiles"          => v"1.0.1"
  "BenchmarkTools"    => v"1.1.0"
  "TSML"              => v"2.7.0"
  "ScikitLearn"       => v"0.6.4"
  "PyCall"            => v"1.92.3"
  "DataFramesMeta"    => v"0.8.0"
  "VegaDatasets"      => v"2.1.1"
  "IJulia"            => v"1.23.2"
  "DecisionTree"      => v"0.10.10"
  "Plots"             => v"1.19.1"
  "RDatasets"         => v"0.7.5"
  "CategoricalArrays" => v"0.9.7"
  "Gadfly"            => v"1.3.3"
  "RCall"             => v"0.13.12"
  "DataFrames"        => v"0.22.7"
```

```julia
using Pkg
Pkg.activate("../../")
Pkg.instantiate()
```

```julia
## How to print

In Julia we usually use `println()` to print
```

```julia
println("I'm excited to learn Julia!")
```

## How to assign variables

All we need is a variable name, value, and an equal's sign!<br>
Julia will figure out types for us.

```julia
my_answer = 42
typeof(my_answer)
```

Output:
```
Int64
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

Output:
```
Float64
```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=

For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

```julia
sum = 3 + 7
```

Output:
```
10
```

```julia
difference = 10 - 3
```

Output:
```
7
```

```julia
product = 20 * 5
```

Output:
```
100
```

```julia
quotient = 100 / 10
```

Output:
```
10.0
```

```julia
power = 10 ^ 2
```

Output:
```
100
```

```julia
modulus = 101 % 2
```

Output:
```
1
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia

```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`

```julia

```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia

```

```julia

```